# Quiz: Hydrograph Fundamentals

This quiz is administered using `nbgrader`. Please answer the following question by assigning the correct letter (`'A'`, `'B'`, `'C'`, or `'D'`) to the variable `answer2` in the code cell below.


### Question 2

**What information can be inferred from the peak of a hydrograph?**

- A. The exact location of the storm’s origin  
- B. The maximum infiltration rate of the soil  
- C. The maximum discharge caused by the storm  
- D. The volume of water stored in the soil

```python
# YOUR ANSWER HERE
answer2 = ''
```

<!-- grade_id: answer2_test, points: 1, grade: true, solution: false -->
```python
# TEST CASE
assert answer2 in ['A', 'B', 'C', 'D'], 'Answer must be one of: A, B, C, or D'
assert answer2 == 'C', 'Incorrect: The correct answer is C – the peak of a hydrograph indicates the maximum discharge.'
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[2], line 2
      1 # TEST CASE
----> 2 assert answer2 in ['A', 'B', 'C', 'D'], 'Answer must be one of: A, B, C, or D'
      3 assert answer2 == 'C', 'Incorrect: The correct answer is C – the peak of a hydrograph indicates the maximum discharge.'

AssertionError: Answer must be one of: A, B, C, or D
```

```python

```